```python
import os 
import pickle
from time import time
```

```python
_version = '7mb'
```

```python
# Replace information
with open('HW2.py', 'r', encoding='utf-8') as file:
    data = file.read()
    
source = "_is_in_develop = True"
output = "_is_in_develop = False"
data = data.replace(source, output, 1)

source = f"_version = '5mb'"
output = f"_version = '{_version}'"
data = data.replace(source, output, 1)

filename = f'HW2-{_version}.py'
with open(filename, 'w', encoding='utf-8') as file:
    file.write(data)
```

```python
if os.path.exists('answer.pickle'):
    !rm answer.pickle

start = time()
!timeout 60 python {filename}
end = time()

# Record your runtime 
duration = end - start

# Read answers and delete temp files
with open('answer.pickle', 'rb') as file:
    answers = pickle.load(file)
!rm answer.pickle
!rm {filename}
```

<!-- grade_id: num-tweets, points: 2, grade: true, solution: false -->
```python
assert answers['num_tweets'] == 1507, "Your answer is not correct"
```

<!-- grade_id: num-unique-users, points: 3, grade: true, solution: false -->
```python
assert answers['num_unique_users'] == 1255, "Your answer is not correct"
```

<!-- grade_id: counts_per_part, points: 5, grade: true, solution: false -->
```python
assert answers['counts_per_part'] == [
    (0, 63),
    (1, 171),
    (2, 24),
    (3, 258),
    (4, 70),
    (5, 242),
    (6, 299),
    (7, 369)
], "Your answer is not correct"
```

<!-- grade_id: num_of_tokens, points: 5, grade: true, solution: false -->
```python
assert answers['num_of_tokens'] == 5997, "Your answer is not correct"
```

<!-- grade_id: num_freq_tokens, points: 5, grade: true, solution: false -->
```python
assert answers['num_freq_tokens'] == 27, "Your answer is not correct"
```

<!-- grade_id: top20, points: 10, grade: true, solution: false -->
```python
assert answers['top20'] == [
    (':', 753),
    ('rt', 666),
    ('.', 538),
    ('the', 422),
    ('trump', 402),
    ('…', 361),
    ('to', 359),
    (',', 344),
    ('in', 272),
    ('is', 258),
    ('a', 255),
    ('!', 203),
    ('of', 201),
    ('for', 196),
    ('and', 182),
    ('on', 157),
    ('i', 155),
    ('he', 138),
    ('that', 132),
    ('"', 124)
], "Your answer is not correct"
```

<!-- grade_id: popular_10_in_each_group, points: 10, grade: true, solution: false -->
```python
assert  answers['popular_10_in_each_group'] == [
    [('bernie', -3.8650704199138914),
     ('to', -4.028408415185754),
     ('that', -4.044394119358453),
     ('"', -4.146841388329271),
     ('and', -4.185866545311334),
     ('-', -4.2223924213364485),
     (',', -4.256339753259786),
     ('i', -4.2761244052742375),
     ('for', -4.292781749227846),
     ('on', -4.294620748891627)],
    [('for', -2.6147098441152083),
     ('that', -2.652076696579693),
     ('bernie', -2.7655347463629774),
     ('rt', -2.840219555963231),
     ('is', -2.841302253980942),
     ('.', -2.8426436720607433),
     (':', -2.898294571920134),
     ('a', -2.9068905956085187),
     ('…', -2.9412661752095333),
     ('"', -2.954196310386875)],
    [('on', -4.972692654004264),
     ('of', -5.066089190457773),
     ('in', -5.08746284125034),
     ('donald', -5.196397212803504),
     ('the', -5.399171093819823),
     (',', -5.426264754702098),
     ('…', -5.49585502688717),
     ('he', -5.523561956057013),
     ('rt', -5.5720234450136585),
     ('.', -5.612030743919327)],
    [('bernie', -1.5025003405291832),
     ('"', -2.4306343543298623),
     ('-', -2.485426827170242),
     ('in', -2.6611980865482416),
     ('will', -2.6865005271832185),
     ('you', -2.6938968722743204),
     ('and', -2.700439718141092),
     ('rt', -2.7211668843194676),
     ('of', -2.7441610955704103),
     ('a', -2.7464259234152726)],
    [('bernie', -3.502500340529183),
     ('for', -3.7078192485066896),
     ('not', -3.7142455176661224),
     ('?', -3.78135971352466),
     ('that', -3.8744691179161412),
     ('to', -3.9642780777660382),
     ('-', -4.0),
     (',', -4.033947331923338),
     ('of', -4.066089190457772),
     ('i', -4.106199403831925)],
    [('not', -1.7142455176661227),
     ('on', -1.9370687442735433),
     ('he', -2.108524456778169),
     ('i', -2.146841388329271),
     ('is', -2.310787537282162),
     ('!', -2.343407822297814),
     ('that', -2.3439544012173616),
     ('.', -2.3849618353734057),
     ('you', -2.3890422907458992),
     ('…', -2.4297658364293984)],
    [('will', -1.40109830832097),
     ('!', -1.4558825515562266),
     ('trump', -1.6968553807920537),
     ('…', -1.7543880404860241),
     ('-', -1.7629608026991508),
     ('in', -1.9377157217456573),
     ('is', -1.9668331360648008),
     ('he', -1.9792414398332026),
     ('rt', -2.1032539617970247),
     ('and', -2.1502426355806126)],
    [('donald', -0.9740047914670558),
     ('trump', -1.6176286896414784),
     ('of', -1.74416109557041),
     ('bernie', -1.8650704199138914),
     ('you', -1.9233787183970876),
     ('the', -1.9397394751825254),
     ('"', -1.9541963103868754),
     (',', -1.9668331360648008),
     ('?', -1.9740047914670555),
     ('i', -1.9907221864119893)]
], "Your answer is not correct"
```