```python
import os 
import pickle
from time import time
import re
```

Input your treasurer's username and an identifier of your program. The cell below is the only cell you need to change.

```python
# leader_name: input your treasurer's username (ucsd.edu email prefix).
leader_name = 'x5zhan'
# Replace 'HW3' to an identifier that lets everyone in your group know this spark application belows to you. 
app_name = 'HW3'
```

<!-- grade_id: sc-inputs, points: 0, grade: true, solution: false -->
```python
# Cell used for grading purposes
```

```python
_HW = 'HW3'
_version = '1gb'
```

```python
# Replace information
with open(f'{_HW}.py', 'r', encoding='utf-8') as file:
    data = file.read()
    
source = r"_is_in_develop = True"
output = f"_is_in_develop = False"
data = re.sub(source, output, data, count=1)

source = r"_version = '.*'"
output = f"_version = '{_version}'"
data = re.sub(source, output, data, count=1)

source = r'sc = set_spark_config(.*)'
output = f'sc = set_spark_config(leader_name="{leader_name}", app_name="{app_name}")'
data = re.sub(source, output, data, count=1)

filename = f'{_HW}-{_version}.py'
with open(filename, 'w', encoding='utf-8') as file:
    file.write(data)
```

- Running this cell will put your Spark application (HW3.py) into your group cluster queue. Your Spark application will be removed from the queue when this cell finishes or the kernel is interrupted/restarted. 

- If your program is stuck in this cell or if you see this message 
    ```
WARN TaskSchedulerImpl: Initial job has not accepted any resources; check your cluster UI to ensure that workers are registered and have sufficient resources
    ```
    it means that your program is in the queue and someone in your group is using the group clusters. Your application will start running automatically when the resources are available (the applications ahead of you have finished). 
    
- To avoid the situation where someone's program is using the resources for too long, all the programs will be timed out after 3600 seconds. **This also means that your program will be timed out if it has been in the queue for 3600 seconds.** 

- If your program finishes after waiting in the queue, your runtime also includes the time of waiting. **In this case, your runtime is highly inaccurate.** You need to read the runtimes statistics printed by your program and manually sum up the runtime of each part. Since the time you waited will be included in the `set up sc` part, you can assume that your `set up sc` part actually took 5 seconds (which is a typical number for setting up the sc). 

```python
with open(f'/datasets/cs255-sp22-a00-public/public/{_version}-runtimes.txt') as file:
    runtimes = list(map(int, file.read().strip().split("\n")))

if os.path.exists('answer.pickle'):
    !rm answer.pickle
    
start = time()
!timeout 3600 python {filename}
end = time()

# Record your runtime 
duration = end - start

!rm {filename}
```

Output:
```
22/04/30 21:09:08 WARN NativeCodeLoader: Unable to load native-hadoop library for your platform... using builtin-java classes where applicable
Using Spark's default log4j profile: org/apache/spark/log4j-defaults.properties
Setting default log level to "WARN".
To adjust logging level use sc.setLogLevel(newLevel). For SparkR, use setLogLevel(newLevel).
[Stage 0:>                                                          (0 + 0) / 2]22/04/30 21:09:25 WARN TaskSchedulerImpl: Initial job has not accepted any resources; check your cluster UI to ensure that workers are registered and have sufficient resources
22/04/30 21:09:40 WARN TaskSchedulerImpl: Initial job has not accepted any resources; check your cluster UI to ensure that workers are registered and have sufficient resources
22/04/30 21:09:55 WARN TaskSchedulerImpl: Initial job has not accepted any resources; check your cluster UI to ensure that workers are registered and have sufficient resources
^C
Traceback (most recent call last):
 
```

<!-- grade_id: correctness, points: 0, grade: true, solution: false -->
```python
correctness = os.path.exists('answer.pickle')
if correctness: 
    with open('answer.pickle', 'rb') as file:
        answers = pickle.load(file)
    with open(f'/datasets/cs255-sp22-a00-public/public/{_version}-answer.pickle', 'rb') as file:
        template = pickle.load(file)

    correctness = answers == template
    assert correctness, "Your answers are not correct. Check the difference between variable answers (your answers) and variable template (correct answers)"
else:
    assert correctness, "Your implementation has runtime error or timed out"
!rm answer.pickle
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-7-ee1c66971076> in <module>
      9     assert correctness, "Your answers are not correct. Check the difference between variable answers (your answers) and variable template (correct answers)"
     10 else:
---> 11     assert correctness, "Your implementation has runtime error or timed out"
     12 get_ipython().system('rm answer.pickle')

AssertionError: Your implementation has runtime error or timed out
```

```python
print(f'Your program took {duration} seconds')
```

<!-- grade_id: runtime-0, points: 5, grade: true, solution: false -->
```python
i = 0
assert correctness and duration < runtimes[i], f'Either your answer is not correct or your runtime is larger than {runtimes[i]} seconds'
```

<!-- grade_id: runtime-1, points: 5, grade: true, solution: false -->
```python
i = 1
assert correctness and duration < runtimes[i], f'Either your answer is not correct or your runtime is larger than {runtimes[i]} seconds'
```

<!-- grade_id: runtime-2, points: 5, grade: true, solution: false -->
```python
i = 2
assert correctness and duration < runtimes[i], f'Either your answer is not correct or your runtime is larger than {runtimes[i]} seconds'
```

<!-- grade_id: runtime-3, points: 5, grade: true, solution: false -->
```python
i = 3
assert correctness and duration < runtimes[i], f'Either your answer is not correct or your runtime is larger than {runtimes[i]} seconds'
```

<!-- grade_id: runtime-4, points: 5, grade: true, solution: false -->
```python
i = 4
assert correctness and duration < runtimes[i], f'Either your answer is not correct or your runtime is larger than {runtimes[i]} seconds'
```

<!-- grade_id: runtime-5, points: 5, grade: true, solution: false -->
```python
i = 5
assert correctness and duration < runtimes[i], f'Either your answer is not correct or your runtime is larger than {runtimes[i]} seconds'
```

<!-- grade_id: runtime-6, points: 5, grade: true, solution: false -->
```python
i = 6
assert correctness and duration < runtimes[i], f'Either your answer is not correct or your runtime is larger than {runtimes[i]} seconds'
```

<!-- grade_id: runtime-7, points: 5, grade: true, solution: false -->
```python
i = 7
assert correctness and duration < runtimes[i], f'Either your answer is not correct or your runtime is larger than {runtimes[i]} seconds'
```

<!-- grade_id: runtime-8, points: 5, grade: true, solution: false -->
```python
i = 8
assert correctness and duration < runtimes[i], f'Either your answer is not correct or your runtime is larger than {runtimes[i]} seconds'
```

<!-- grade_id: runtime-9, points: 5, grade: true, solution: false -->
```python
i = 9
assert correctness and duration < runtimes[i], f'Either your answer is not correct or your runtime is larger than {runtimes[i]} seconds'
```